# Homework 5: Iterators, Memoization and Exceptions(18 pts)

name: Mingyang Zhou 

email: zmyickle@umich.edu

This homework assignment took me 1 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

---

I did not discuss this homework with anyone

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

---
##  Problem 1: Iterator (6 pts)
Implement a custom range iterator that mimics some of the basic functionality of Python's built-in `range()` function (You can not use `range()` direclty). This exercise will help you understand how iterators work in Python.

**Note:** there is the same question as our in-class practice.

##  Problem 1.1: A custom class called `MyRange` (3 pts)
1. Create a class named `MyRange` that takes two integer arguments in its constructor:

     **start:**  the first value of the range

   
     **start:**: the value to stop before (exclusive)

3. The class should implement the iterator protocol, making it usable in a for loop.

4. Each iteration should return the next integer in the sequence, starting from start and incrementing by 1 each time.
The iteration should stop when the value would become equal to or greater than end.

<!-- grade_id: class_my_range, points: 0, grade: false, solution: true -->
```python
class MyRange:
    # YOUR CODE HERE
    def __init__(self,start,end):
        self.start = start
        self.end = end
        self.current = start

    def __iter__(self):
        return self

    def __next__(self):
        if self.current < self.end:
            next_value = self.current
            self.current +=1
            return next_value
        else:
            raise StopIteration

    
```

<!-- grade_id: correct_MYRANGE, points: 3, grade: true, solution: false -->
```python
nums = MyRange(1, 5)
assert next(nums) == 1, "First value should be 1"
assert next(nums) == 2, "Second value should be 2"
assert next(nums) == 3, "Third value should be 3"
assert next(nums) == 4, "Fourth value should be 4"

try:
    next(nums)
    assert False, "StopIteration should have been raised"
except StopIteration:
    pass

```

##  Problem 1.2: A custom generator called `my_range` (3 pts)
Create a function `my_range(start, end)` that generates a sequence of numbers from start to end (exclusive). The function should be implemented as a generator.

<!-- grade_id: class_my_range2, points: 0, grade: false, solution: true -->
```python
def my_range(start, end):
    # YOUR CODE HERE
    current = start
    while current < end:
        yield current
        current +=1

```

<!-- grade_id: correct_range, points: 3, grade: true, solution: false -->
```python
nums = my_range(1, 5)
assert next(nums) == 1, "First value should be 1"
assert next(nums) == 2, "Second value should be 2"
assert next(nums) == 3, "Third value should be 3"
assert next(nums) == 4, "Fourth value should be 4"

try:
    next(nums)
    assert False, "StopIteration should have been raised"
except StopIteration:
    pass

```

##  Problem 2: Memoization (6 pts)

###  2.1: Climbing stairs (3 pts)
You are climbing a staircase. It takes n steps to reach the top.

Each time you can either climb 1 or 2 steps. In how many distinct ways can you climb to the top? Let's try not use recursion in your implementation. Instead, use an iterative approach or use memoization.

<!-- grade_id: climbStairs, points: 0, grade: false, solution: true -->
```python
def climb_stairs(n):
    # YOUR CODE HERE
    if n==0:
        return 1
    elif n==1:
        return 1

    first = 1 
    second = 1 

    for i in range(2, n + 1):
        current = first + second  
        first = second  
        second = current

    return second
```

<!-- grade_id: correct_stairs, points: 3, grade: true, solution: false -->
```python
assert climb_stairs(1) == 1, "Test case 1 failed"
assert climb_stairs(2) == 2, "Test case 2 failed"
assert climb_stairs(3) == 3, "Test case 3 failed"
assert climb_stairs(7) == 21, "Test case 4 failed"
assert climb_stairs(35) == 14930352, "Test case 5 failed (larger input)"
```

###  2.2: Two Sum(3 pts)
Given an array of integers nums and an integer `target`, return indices of the two numbers such that they add up to `target`.

You may assume that each input would have exactly one solution, and you may not use the same element twice. You can return the answer in any order.

**Note:**: This is a classical interview question. Try to use a dictionary as a "mini database" to store and quickly look up values.

<!-- grade_id: sum, points: 0, grade: false, solution: true -->
```python
def two_sum(nums, target):
    # YOUR CODE HERE
    num_dict={}

    for index, num in enumerate(nums):
        another = target - num 

        if another in num_dict:
            return [num_dict[another],index]
        
        num_dict[num] = index

    return []
```

<!-- grade_id: correct_two_sum, points: 3, grade: true, solution: false -->
```python
assert two_sum([2, 7, 11, 15], 9) == [0, 1], "Test case 1 failed"
assert two_sum([3, 2, 4], 6) == [1, 2], "Test case 2 failed"
```

##  Problem 3: Error Handling (6 pts)

###  3.1: Safely Divide(3 pts)
Write a function called `safe_divide(a, b)` that safely performs division of a by b. The function should:
1) Return the result of a divided by b if b is not zero.

2) Return the string "Cannot divide by zero" if b is zero.

3) Use a try-except block to handle the `ZeroDivisionError`.


<!-- grade_id: divide, points: 0, grade: false, solution: true -->
```python
def safe_divide(a, b):
    # YOUR CODE HERE
    try:
        ans = a/b
        return ans
    except:
        ZeroDivisionError
        return "Cannot divide by zero"

```

<!-- grade_id: correct_divide, points: 3, grade: true, solution: false -->
```python
assert safe_divide(10, 2) == 5.0, "Test case 1 failed"
assert safe_divide(10, 0) == "Cannot divide by zero", "Test case 2 failed"
assert isinstance(safe_divide(10, 3), float), "Test case 5 failed: Result should be a float"
```

###  3.2: Raise our own error (3 pts)
Create a function called `safe_divide(a, b)` that does the following:

Takes an age parameter (integer).
1) If age is not int or float, raise `TypeError` with message "Age must be a number".
2) If the age is less than 0, raise a `ValueError` with message "Age cannot be negative".
3) If the age is over 150, raise the `ValueError` with the message "Age cannot be over 150".



<!-- grade_id: validate, points: 0, grade: false, solution: true -->
```python
def validate_age(age):
    # YOUR CODE HERE
    if not isinstance(age,(int,float)):
        raise TypeError("Age must be a number")

    if age < 0:
        raise ValueError("Age cannot be negative")
    
    if age > 150:
        raise ValueError("Age cannot be over 150")

    return f"Valid age: {age}"
```

<!-- grade_id: correct_error, points: 3, grade: true, solution: false -->
```python
assert validate_age(25) == "Valid age: 25", "Test for valid age failed"
try:
    validate_age("twenty")
    assert False, "Non-numeric age should raise TypeError"
except TypeError as e:
    assert str(e) == "Age must be a number", "Incorrect error message for non-numeric input"
    
try:
    validate_age(-5)
    assert False, "Negative age should raise ValueError"
except ValueError as e:
    assert str(e) == "Age cannot be negative", "Incorrect error message for negative age"

```